```python
# Initialize Otter
import otter
grader = otter.Notebook("lab11.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 11 – Dictionaries

## Introduction to Computational Thinking with Data Science and Society
 Welcome to Lab 11! This week we will be covering dictionaries. While other data types we've seen in class are quite useful in many ways, dictionaries have a special purpose.

```python
# Just run this cell to load in the relevant dependencies

from datascience import *
import numpy as np
import matplotlib.pyplot as plt
from pathlib import Path
from ipywidgets import interact, widgets
from IPython.display import HTML, display
%matplotlib inline
plt.style.use('ggplot')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## [Tutorial] Part 1: Dictionaries


<img src='dictionary.png' width=300>

**Dictionaries** can be very useful. They store key/value pairs that can be used to map one value to another. You can think of a dictionary as a list where the indexes (locations) of the values of the list are no longer their integer locations, but rather their keys.

>In an array, you access the first item with `my_array.item(0)`.

>In a dictionary, you access the "key-th" item with `my_dictionary[key]`.

If we think of list items as having their "address" be their location in the list, then a dictionary value's "address" is its key.

Some important properties of dictionaries to note:
- The key and value **do not** have to be of the same type
- We designate a new key/value entry in a dictionary in this format: *key* **:** *value*
- We store all these key/value entries in a dictionaries with braces `{}` around the ends (like `[]` with a list) and commas separating the entries
- Keys in a dictionary are unique, but values don't have to be unique. 
    - For example {'a': 100, 'a': 200} is not a valid dictionary
    - but, {'a': 100, 'b': 100} is a valid dictionary
    
Let's take a closer look at a dictionary in practice:

```python
my_dictionary = {"a": 100, "b": 200, "c": 300}
print("The value 'a' maps to the value:", my_dictionary["a"])
print("The value 'b' maps to the value:", my_dictionary["b"])
print("The value 'c' maps to the value:", my_dictionary["c"])
```

### How to Access the Data

We can't access a dictionary's values like we can access a list's values. If we want the "first" item in a dictionary, we cannot ask for `my_dictionary[0]`, because this request is really asking "What does the key 0 map to in this dictionary?". If your dictionary does not have a value associated with the key 0, you will get an error.

```python
my_dictionary[0]
```

A `KeyError` warning means that you asked for a key that is not in your dictionary. This may happen when you are writing a function with a dictionary, so if you see it, this is what it means.

### Changing the Data

We can add the key value pair `(key, value)` with the following syntax:

> `my_dictionary[key] = value`

Run the cell below to change the `"d"` entry of the dictionary to 400:

```python
my_dictionary["d"] = 400 # Add the key/value pair ("d", 400) to our dictionary
my_dictionary
```

We can use **any** data type we know as a value in a dictionary...

```python
# Here, the value we add is a list!
my_dictionary["grocery list"] = make_array("apples", "bananas", "carrots")
my_dictionary
```

...including even having a **dictionary itself** as a value! 

```python
my_dictionary["squares"] = {1: 1, 2: 4, 3: 9, 4: 16}
my_dictionary
```

### Dictionary Iteration

We can get a list of a dictionary's keys with the `.keys()` function.

```python
my_keys = my_dictionary.keys()
my_keys
```

```python
# Note the type of this list of keys
type(my_keys)
```

To iterate over the keys in a dictionary, we can use a `for` loop!

```python
for key in my_dictionary:
    print("I am a key, and my name is:", key)
```

We can also get a list of a dictionary's values with the `.values()` function.

```python
my_values = my_dictionary.values()
my_values
```

We can iterate over the values of a dictionary like this:

```python
for value in list(my_dictionary.values()):
    print("I am a value, and my name is:", value)
```

We can use this to do cool things like change all the values in a dictionary!

```python
def add_one_to_dictionary_values(dictionary):
    for key in dictionary:
        dictionary[key] = dictionary[key] + 1
    return dictionary

new_dictionary = {"data": 6, "cs": 61, "poli sci": 1}
modified_dictionary = add_one_to_dictionary_values(new_dictionary)
modified_dictionary
```

We can also use the `.items()` function to return all key-value pairs in a dictionary as a list-like object:

```python
# just run this cell
for k, v in modified_dictionary.items():
    print(k, v)
```

In the cell above, the names `k` and `v` are respectively assigned to each key and value in the `modified_dictionary` dictionary. If you’re curious about how this works, please ask us!

A final note: When provided a dictionary name, Python by default will iterate over the keys of the dictionary:

```python
# just run this cell
for k in modified_dictionary:
	print(k)
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part 2: Dictionary Fundamentals



## Question 1 – Emojify 

The default keyboard on iOS suggests emojis for you to use in place of boring, ordinary words.

<img src = "https://support.apple.com/library/content/dam/edam/applecare/images/en_US/iOS/ios12-iphone-x-messages-replace-words-with-emoji.jpg" width=200>

In this question, you will replicate some of that behavior using dictionaries!


### Emojis in Python
In Python, emojis can be included as part of a string. For example:

```python
'🤤'
```

If you remove the quotes from the emoji above, you will see `SyntaxError: invalid character in identifier`. **Make sure that throughout this question, your emojis are contained within strings!** <br> Fun fact: they cannot currently be used as variable names. Try it and see what error you get!

---
## Question 1a

In the cell below, define a dictionary `fav_emojis` that has the following **five** keys:
- `'happy'`
- `'annoyed'`
- `'tired'`
- `'love'`
- `'food'`

The values corresponding to these five keys must be an emoji. [getemoji.com](https://getemoji.com) allows you to copy and paste emojis. To select an emoji, double click it to highlight it. You may choose any emojis you would like **as long as**:
>- it is copied from the site above
>- it is not in the "New Emojis" category at the bottom

Have fun with it! We've chosen an emoji for `'happy'` for you, but feel free to change it.

**Some troubleshooting tips:**
- After defining your dictionary, you may see some emoijs displayed with `'\U001...'` instead of their actual graphic. When a browser doesn't support the emoji's encoding, the code is displayed instead. If this happens, pick different emojis.
- If you fail the test case that says your emojis are invalid, and you're certain you correctly defined your dictionary, you may consider choosing other emojis that are more generic that are more likely to be recognized by our autograder. This most likely won't be a problem.


```python
fav_emojis = {
    'happy': '😁',
    ...
}

fav_emojis
```

```python
grader.check("q1a")
```

---
## Question 1b

Now, complete the implementation of the function `emojify`, which takes in a string `message` and returns a new string with all instances of any of the keys in `fav_emojis` replaced with their corresponding emoji value. Example behavior is shown below, though the emojis will be different, depending on what you put in `fav_emojis`. If you passed the previous question, you don't need to change your emojis!

```py
>>> emojify('Filing taxes makes me tired and want food.')
'Filing taxes makes me 😵 and want 🌽.'

>>> emojify('I LOVE love life right now. I am so happy – why do you look so annoyed?!')
'I 💋 💋 life right now. I am so 😁 – why do you look so 💀?!'

>>> emojify("It's not you, it's me... I don't make you haPPy, I make you tired.")
"It's not you, it's me... I don't make you 😁, I make you 😵."
```


```python
def emojify(message):
    # This line ensures your code replaces correctly if any of
    # the keys in fav_emojis appears in uppercase in the message
    message = message.lower()

    ...
    
    # Don't change this
    return message
```

```python
grader.check("q1b")
```

### Fun Demo

Run the cell below to produce a text box (don't worry about the code itself). Type text in the text box and watch it get emojified live!

```python
def emojify_live(type_here):
    display(HTML('<h2>' + emojify(type_here) + '</h2>'))
interact(emojify_live, type_here="I LOVE food");
```

<!-- BEGIN QUESTION -->

---
## Question 1c

Nice and simple: What's your favorite emoji? Place it in the Markdown cell below.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br></br>
<hr style="border: 1px solid #fdb515;" />

## Question 2 – JSON to Python Dictionary 


---
## Question 2a

We can read and convert JSON files into Python dictionaries. That's what you'll do in this question.

Before following these instructions, make sure to save your notebook (which you should be doing frequently anyways)!

1. Right click the Jupyter logo in the top left of your screen, and click "Open Link in New Tab" (it may appear as Open...)
2. Navigate to the lab11 folder
3. Identify the name of the `.json` file that contains Google Maps data. You may have to open both `.json` files to determine which one it is; you can open files by clicking on them.
4. Set the string `maps_path` below equal to the path to that file. `maps_path` end with `'.json'`.


```python
maps_path = ...
```

```python
grader.check("q2a")
```

If you answered the previous part correctly, you should be able to run the following cell:

```python
# just run this cell
    
# we defined this function for you
def read_json2(path):
    import json
    with open(path, 'r') as f:
        contents = json.load(f)
    return contents


maps_data = read_json2(maps_path)
maps_data

```

---
## Question 2b

Next, assign `key_1`, `key_2`, and `key_3` below so that `maps_data[key_1][key_2][key_3]` evaluates to the latitude of the location whose data is stored in `maps_data`. We've done `key_2` for you.

_Hint_: Work one step at a time. You know that `key_1` must be one of the six keys in `maps_data_keys`, which you found above. Then, given what we've set `key_2` to, what must `key_3` be?


```python
key_1 = ...
key_2 = 'location'
key_3 = ...

maps_data[key_1][key_2][key_3]
```

```python
grader.check("q2b")
```

---
## Question 2c

The dictionary above is quite unwieldy, and contains many nested dictionaries! Let's try and extract some data from it programatically (that is, using code).

Write a `for` loop that prints out each key and the data type of each value. Complete the code below so that it prints the below lines, possibly not in the same order:

```
address_components <class 'list'>
formatted_address <class 'str'>
geometry <class 'dict'>
place_id <class 'str'>
plus_code <class 'dict'>
types <class 'list'>

```
_Hints_:
- See the `for` loop access pattern in the tutorial before Q1. Which of the three dictionary methods should you use?
- Recall that `type(x)` returns the data type of `x`.




```python
for ...:
    ...


```

```python
grader.check("q2c")
```

By the way, `maps_data` contains location information for 84 Viet, a Vietnamese restaurant in Downtown Berkeley. It's quite good, you should try it!

## Pet of the Day

Congratulations on completing Lab 11! Can you help me again with these weights?

<img src="cat-weights.jpg" width="50%">

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1a (points: 2)
```
>>> len(fav_emojis) == 5
True
```
```
>>> sorted(list(fav_emojis.keys())) == ['annoyed', 'food', 'happy', 'love', 'tired']
True
```
```
>>> all([len(e) == 1 for e in fav_emojis.values()])
True
```
```
>>> def read_json(path):
...     import json
...     with open(path, 'r') as f:
...         contents = json.load(f)
...     return contents
>>> all_emojis = read_json('emojis.json')
>>> invalid_keys = []
>>> for k in fav_emojis.keys():
...     if fav_emojis[k] not in all_emojis:
...         invalid_keys.append(k)
>>> len(invalid_keys) == 0
True
```


Tests for q1b (points: 4)
```
>>> convert_1 = emojify('Filing taxes makes me tired and want food.').split(' ')
>>> e1_1, e1_2 = (convert_1[4], convert_1[-1][:-1])
>>> e1_1 == fav_emojis['tired'] and e1_2 == fav_emojis['food']
True
```
```
>>> convert_2 = emojify('I LOVE love life right now. I am so happy – why do you look so annoyed?!').split(' ')
>>> e2_1, e2_2, e2_3, e2_4 = (convert_2[1], convert_2[2], convert_2[9], convert_2[-1][:-2])
>>> e2_1 == fav_emojis['love'] and e2_2 == fav_emojis['love'] and (e2_3 == fav_emojis['happy']) and (e2_4 == fav_emojis['annoyed'])
True
```
```
>>> convert_3 = emojify("It's not you, it's me... I don't make you haPPy, I make you tired.").split(' ')
>>> e3_1, e3_2 = (convert_3[9][:-1], convert_3[-1][:-1])
>>> e3_1 == fav_emojis['happy'] and e3_2 == fav_emojis['tired']
True
```
```
>>> convert_4 = emojify('Billy is lovesick, and not the kind that can be solved with food – Billy is annoyed, really.').split(' ')
>>> e4_1, e4_2, e4_3 = (convert_4[2][0], convert_4[12], convert_4[-2][0])
>>> e4_1 == fav_emojis['love'] and e4_2 == fav_emojis['food'] and (e4_3 == fav_emojis['annoyed'])
True
```


Tests for q2a (points: 1)
```
>>> isinstance(maps_path, str)
True
```
```
>>> maps_path == '84_viet.json'
True
```


Tests for q2b (points: 2)
```
>>> isinstance(key_1, str) and isinstance(key_2, str)
True
```
```
>>> key_1 == 'geometry' and key_3 == 'lat'
True
```
```
>>> np.isclose(maps_data[key_1][key_2][key_3], 37.8727848)
True
```


Tests for q2c (points: 1)
```
>>> solution = maps_data
>>> maps_data.items() == solution.items()
True
```
